```python
# Initialize Otter
import otter
grader = otter.Notebook("hw04.ipynb")
```

# Homework 4 

## Due Date: Wednesday, Sept. 10th by 11:59 PM MT on Gradescope


## Assignment Learning Objectives
 - Calculate probabilities for equally likely outcomes
 - State and apply the Inclusion/Exclusion Principle
 - Translate given information about events into a joint probability table
 - State the mathematical definition of conditional probability and use it to  calculate conditional probabilities.
 - State the multiplication rule and use it to calculate joint probabilities.
- Use Numpy to simulate probabilities
- Define and apply the Law of Total Probability
- Define and apply Bayes' Theorem to update probabilities


## Collaboration and Generative AI Policy

Data science is collaborative, but a key part of learning is **doing the intellectual work yourself**. Homework is not just about getting the correct answer—it is practice for developing the reasoning and conceptual understanding you will need on quizzes and exams.

You may discuss concepts, ask for help when you are stuck, and use approved references, but the solutions you submit must be your own. **You should be able to independently reproduce, explain, and justify any work you submit if asked by an instructor or TA.** More importantly, you should understand the underlying concepts well enough to recognize when they apply and use them in a new problem or context.

### If You Choose to Use HiTA (https://www.hita.ai/)

HiTA is an optional GenAI tutoring tool available specifically for this course. It is permitted because it is designed to support learning rather than replace it. HiTA provides explanations, hints, questions, and guidance while leaving the substantive reasoning and problem-solving to you.

HiTA is the only permitted Gen-AI tool in this course. ChatGPT, Gemini, Claude, Copilot, and **other generative AI tools may not be used to obtain assistance on course assignments**.

If you choose to use HiTA, treat it as a **last resort for help, not a first step toward a solution**. Homework is an opportunity to practice the conceptual reasoning and problem-solving skills you will need on quizzes and exams, where GenAI will not be available.

Use the following guidelines:

1. **Try the problem yourself first—and give yourself time to struggle with it.** Productive struggle is an important part of learning. Before asking HiTA for help, work on the problem long enough to identify what you understand, what you do not understand, and exactly where your reasoning breaks down. If possible, give yourself some time away from the problem before turning to HiTA; continuing to think about a difficult concept, even when you are not actively working on it, is part of developing the problem-solving skills you will need on exams.

2. **Ask for guidance, not answers.** If you are still stuck, ask HiTA for a hint, conceptual explanation, or help understanding why an approach might or might not work. Do not ask it to select the method, carry out the substantive reasoning, or produce the complete solution for you.

3. **Focus on why, not just how.** Getting the correct homework answer is not the goal—the goal is learning the underlying concepts. Being able to reproduce an algorithm or sequence of calculations does not necessarily mean you understand it. You should be able to explain why a method works, when it applies, when it does not apply, and how to use the same underlying idea in a substantially different problem or application.

4. **Evaluate HiTA critically.** HiTA can make mistakes. If you already have a solution that you believe is correct, you may use HiTA to critique or check your reasoning, but do not assume its response is correct simply because it sounds convincing. GenAI can hallucinate or provide incorrect reasoning, so if HiTA's response conflicts with your understanding or leaves you uncertain, check with Course Staff on Piazza or during Office Hours before relying on it.

5. **Return to the problem independently.** After receiving help from HiTA, put its response aside and work through the problem yourself. You should be able to reconstruct and explain the reasoning without referring back to HiTA. If you cannot, you have not yet learned the material well enough and should continue working on the concept or seek help from Course Staff.

A good test is whether using HiTA ultimately makes you **better able to solve problems without HiTA**. If you become dependent on it whenever you are uncertain, you are practicing the opposite of the skill you will need on quizzes, exams, and in later courses.

The key principle is to **avoid cognitive offloading**: do not transfer to HiTA the thinking and reasoning that you need to learn to do yourself.

#### Homework and Project Collaboration: Permitted and Not Permitted

The following are **violations** of the CSCI 3022 Academic Honesty and Collaboration Policy:

- Consulting a webpage or other source that contains a solution to the assigned problem or a similar problem.
- Reading another student's solution before writing your own.
- Copying another student's work or an online solution and then editing or modifying it.
- Copying three or more lines of code or mathematical work from another student, or allowing another student to copy three or more lines of your work.
- Asking someone else to write all or part of a program or solution for you, or to provide the code needed to fix an error except for simple syntax errors.
- Having another person or a GenAI tool convert an entire solution, or a substantial portion of one, into LaTeX for you.
- Copying, pasting, or uploading assignment text or screenshots into any GenAI tool other than HiTA.
- Searching Google or asking a GenAI tool other than HiTA an identical or substantially similar question to one appearing on an assignment.
- Using any GenAI tool other than HiTA to solve, check, rewrite, or substantially complete an assignment problem.
- Using HiTA to obtain a complete solution rather than an explanation, hint, or guidance.
- Submitting an exact or lightly edited solution generated by an LLM, including HiTA.
- Falsifying evidence that you submitted an assignment,

The following are **permitted**:

- Working on a problem independently first, then discussing a specific point where you are stuck and returning to complete your solution independently.
- Looking up individual programming or LaTeX commands and syntax.
- Asking how a programming or LaTeX construct works or why it is not behaving as expected.
- Asking someone to help you find or understand an error in your program without providing the solution for you.
- Consulting course notes, the textbook, or references that provide syntax, formulas, or conceptual explanations without providing a solution to the assigned problem.
- Using HiTA for permitted forms of assistance, such as conceptual explanations, hints, clarification, or guidance that helps you develop and understand your own solution.

The key distinction is whether a resource is helping you learn how to do the work yourself or is doing part of the intellectual work for you. Your submitted work must be your own, and you should be able to independently explain and reproduce it if asked. When in doubt, rely on the course notes, course textbook, and references that provide syntax and/or formulas, or ask the instructional staff.



#### Grading
Grading is broken down into autograded answers and manually graded answers. 

For autograded answers, the results of your code are compared to provided and/or hidden tests.

For manually graded answers you must show and explain all steps.  Graders will evaluate how well you answered the question and/or fulfilled the requirements of the question.


### Score breakdown



Question | Points | Grading Type
--- | --- | ---
Question 1 | 4 | autograded
Question 2 | 4 | manual
Question 3 | 8 | autograded
Question 4a | 4 | manual
Question 4bi-ii | 5 | autograded
Question 4biii | 4 | manual 
Question 5 | 8 | manual
Question 6 | 7 | manual
Question 7 | 6 | manual
|Total | 50

## Recommended Readings for this HW:

- [Probability of Equally Likely Outcomes](https://chrispiech.github.io/probabilityForComputerScientists/en/part1/equally_likely/)
- [Law of Total Probability](https://chrispiech.github.io/probabilityForComputerScientists/en/part1/law_total/)
- [Bayes Theorem](https://chrispiech.github.io/probabilityForComputerScientists/en/part1/bayes_theorem/)

- [Computational and Inferential Thinking: Section 9.5 - Finding Probabilities](https://inferentialthinking.com/chapters/09/5/Finding_Probabilities.html)
- [Computational and Inferential Thinking: Chapter 18: Updating Predictions](https://inferentialthinking.com/chapters/18/Updating_Predictions.html)
- [Computational and Inferential Thinking: Section 9.3 - Simulation](https://inferentialthinking.com/chapters/09/3/simulation/)
- [Theory Meets Data: Chapters 1, 2 and Sections 3.1 and 3.2](https://data88s.org/textbook/content/intro.html) 





<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />



In this homework you'll investigate probability. You'll also use probability simulations to estimate probabilities, gain intuition about random processes and to check your pencil and paper work.   


```python
import pandas as pd
import numpy as np
import hashlib
import matplotlib.pylab as plt
%matplotlib inline


def get_hash(num):
    """Helper function for assessing correctness"""
    return hashlib.md5(str(num).encode()).hexdigest()
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴  Question 1: Roulette ##


A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner. 

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. Below is an example of a roulette wheel!

<img src="img/roulette_wheel.jpeg" width="330px">


On each spin of a roulette wheel, all 38 pockets are equally likely to be the winner regardless of the results of other spins. Among the 38 pockets, 18 are red, 18 black, and 2 green. 

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 1a (1 pt) ## 
    
**In each part below, write code that evaluates to the chance of the event described.**
    
What is the probability the winning pocket is black on all of the first **four** spins?


```python
first_three_black = ...
```

```python
grader.check("q1a")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 1b (1 pt) ## 

What is the probability the color green never wins in the first 8 spins?

```python
no_green = ...
```

```python
grader.check("q1b")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 1c (1 pt) ## 

What is the probability the color green wins at least once on the first 8 spins?


```python
at_least_one_green = ...
```

```python
grader.check("q1c")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 1d (1 pt) ## 


What is the probability two of the three colors never win in the first 8 spins?



```python
lone_winners = ...
```

```python
grader.check("q1d")
```

<!-- BEGIN QUESTION -->

***


<a id='p2'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 2 (4 pts) ## 

Let $A$ and $B$ be events in a sample space $S$.  

Suppose that the probability that $A$ occurs is $0.3$, the probability that $B$ occurs is $0.4$, and the probability that neither $A$ nor $B$ occur is $0.5$ (that is, the joint probability that $A$ doesn't occur AND $B$ doesn't occur is  $0.5$). 

**Determine if you have enough information to answer the following questions.  If so, write your solution showing all steps.  If not, explain what additional information you would need.**


**2a) (2 pts)**  What is $P(A, B)$?

**2b) (2 pts)**  What is $P(A \mid B^C)$?

**Hint:  Use the information provided to create a joint probability table for $A$ and $B$.**

Note: You can easily create a table in Markdown using this format (double click on this cell to see the syntax to create this table): 
|   | Col A Title|Col B Title|
|---|---|---|
|**Row A Title**|Entry 1  |other stuff  |
|**Row B Title**|Latex works: $\int_1^2 x\,dx$ |  more stuff |


Write up your full solution to both questions in the SAME box below using LaTeX (not code).  Show all steps fully justifying your answers.  


_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for Problem 2a, enter the answer (just the probability) you came up with in the cell below.  If you said you didn't have enough information, enter the number 0.  Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the cell above.  

```python
q2a_answer = ...
```

```python
grader.check("q2acheck")
```

To check your answer for Problem 2b, enter the answer you came up with (just the probability) in the cell below. If you said you didn't have enough information, enter the number 0. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX fully justifying your answer in the Markdown cell provided above.  

```python
q2b_answer = ...
```

```python
grader.check("q2bcheck")
```

***



<a id='p3'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 3: Simulating Probabilities Video Assignment (8 pts) ## 

#### Complete this [Simulating Probabilities Video Assignment (click this link)](https://canvas.colorado.edu/courses/141915/assignments/2764452) available on Canvas (to receive points you must watch the entire video and answer the embedded questions).


As you watch the video, fill in the code below along with the video so you can practice with the concepts in the video.

## Simulating Probabilities

As we discussed in lecture, the frequentist definition of probability defines it in terms of a limit:


$$P(Event) = \lim_{n \to \infty} \frac{count(Event)}{n}$$



In English this reads: Perform n trials of an "experiment" which could result in a particular “Event” occurring. 

The probability of the event occurring, P(Event) , is the ratio of trials that result in the event, written as  count(Event), to the number of trials performed, n. 

In the limit, as your number of trials approaches infinity, the ratio will converge to the true probability.


One of the big payoffs of simulation is that it can let us answer some probability questions that are otherwise quite difficult.  We can instead just **simulate the experiment a large number of times and get approximate results based on simulation.**


There are several ways to draw random samples using Python:


#### Random Sampling in Pandas: 
- `df.sample(n)` (https://pandas.pydata.org/docs/reference/api/pandas.DataFrame.sample.html) draws a random sample of `n` rows from the DataFrame `df`. The output is a DataFrame consisting of the sampled rows. 



#### Random Sampling in Numpy:  https://numpy.org/doc/stable/reference/random/index.html

The `numpy.random`module implements pseudo-random number generators (PRNGs or RNGs, for short) with the ability to draw samples from a variety of probability distributions.   In general, you will create a Generator instance with default_rng and call the various methods on it to obtain samples from different distributions.


```python
# Create a Generator instance:
rng = np.random.default_rng() 
```


- `rng.choice(a)` (https://numpy.org/doc/stable/reference/random/generated/numpy.random.choice.html) draws a random sample from a population whose elements are in an array `a`. The output is an array consisting of the sampled elements.

- `rng.shuffle(a)`
(https://numpy.org/doc/stable/reference/random/generated/numpy.random.Generator.shuffle.html#numpy.random.Generator.shuffle)  Modifies an array or sequence in-place by shuffling its contents. 
  
Warning: The pseudo-random number generators implemented in the `numpy.random` module are designed for statistical modeling and simulation. They are not suitable for security or cryptographic purposes. See the `secrets` module from the standard library for such use cases.




### Estimating Simple Probabilities 
*** 

Suppose we wanted to randomly sample from the first 6 letters in the alphabet:

```python
select_from =["A", "B", "C", "D", "E", "F"]
```

We can simulate randomly selecting from this list using `rng.choice`, which returns a randomly selected entry from a Numpy array.  
If no optional parameters are passed in, `rng.choice` assigns an **equal probability** to each entry of the array.   

```python
rng.choice(select_from)
```

We can simulate repeatedly selecting (with replacement) by specifying the size of an array:

```python
#Write code to randomly sample 10 times WITH replacement
...
```

```python
#Write code to randomly sample 5 times WITHOUT replacement
...
```

**Simulating Coin Tosses**

As a simple example, consider a fair coin.  We can represent the sample space for this coin with a Numpy array with two entries: "H" and "T"

```python
coin = np.array(["H","T"])
coin
```

```python
rng.choice(coin)
```

We can simulate many flips of the coin and store the results in an array by passing the size parameter:

```python
flips = rng.choice(coin, size=10)
print(flips)
```

We can also count the number of times we get heads:

```python
flips = rng.choice(coin, size=10)
print(flips)
sum(flips == 'H')
```

We can alter the probability with which `choice` selects a particular entry of the sample space array by passing in an optional array of probabilities, e.g. p = [0.75, 0.25]:

```python

biased_flips = rng.choice(["H", "T"], p=[0.75, 0.25], size=10)

print(biased_flips)
sum(biased_flips == 'H')
```

*** 
###  Example 1: 
Simulate rolling a 6-sided die, 10 times:

```python
die = ...
roll = ...
roll
```

**SideNote**: If instead of passing a list/array you pass an int, to `rng.choice` the random sample is generated as if it were `np.arange(a)`

```python
np.arange(1,10,2)
```

```python
rng.choice(5, size=10)
```

*** 
### Example 2- Flipping a Fair Coin

Now suppose we want to run a simple simulation to estimate the probability  that the coin comes up Heads (which we expect to be $0.5$ because the coin is fair).  One way to do this is to do a large number of coin flips and then divide the number of flips that come up Heads by the total number of flips. The following code flips the coin 50 times and computes the desired ratio: 

```python
# The seed() method is used to initialize the random number generator.
# The random number generator needs a number to start with (a seed value),
#   to be able to generate a random number.
# By default the random number generator uses the current system time.
# We can use the seed() method to customize the start number of the random number generator
# so that we get the same results each time.

rng = np.random.default_rng(seed=152)


# Write code to simulate flipping a coin 50 times and approximating the probability of heads using your simulation:
flips = ...
approx_prob_heads = ...

print("the probability of heads is approximately {:.3f}".format(approx_prob_heads))
```

OK, so the simulation estimated that the probability of the coin coming up heads is pretty far off from the $0.5$ that we expected.  This is likely because we didn't do very many coin flips.  Let's see what happens if we rerun the simulation with $500$ coin flips. 

```python

flips = rng.choice(coin, size=50000)
approx_prob_heads = sum(flips=="H")/len(flips)
print("the probability of heads is approximately {:.3f}".format(approx_prob_heads))
```

With $500$ coin flips our estimate came out much closer to $0.5$



It's an interesting exercise to make a plot of the running estimate of the probability as the number of coin flips increases.    This is one way to visualize a simulation of the frequentist definition of probability. 


We'll use the same random sequence of coin flips from the previous simulation.  

We'll be using the Python `range` function (https://docs.python.org/3/library/functions.html#func-range) in the following code, so let's see how it works:


```python
# Notice how range function works:

for ii in range(5):
    print(ii)
```

```python

def plot_estimates(n):
    """ Simulates flipping a fair coin n times and returns running \\
        estimate of the probability of getting a heads as num_trials gets larger"""

        # Ex: Running_prob is a list with ratio of 'H' being counted per times coin flipped.
        # Suppose we got T, H, T, ...
        # running_prob[0]=0/1   running_prob[1]=1/2   running_prob[2]=1/3   etc
        # running_prob = [0, 0.5, 0.3333..., etc]
    
    flips = rng.choice(["H","T"], size=n)
    running_prob = []

    for ii in range(n):
        num= ...
        denom= ...
        running_prob.append(num/denom)
         
    
    return running_prob

# Run code for num trials
num_trials=10000

p = ...
x = ...

print("the probability of heads is approximately {:.3f}".format(p[num_trials-1]))
 
# Plot running estimate of probability of getting heads as num_trials gets larger:
fig, ax = plt.subplots(figsize=(12,6))

# plot the terms in p
ax.plot(..., ...  , color="steelblue",label="Simulated probability")

#Plot the theoretical probability
plt.axhline(y = ..., color = 'r', linestyle = '-', label = "Theoretical probability")

# put labels on the axes and give the graph a title.
ax.set_title("Running Estimate of Probability of Heads", fontsize=20)
ax.set_xlabel("Number of Flips", fontsize=16)
ax.set_ylabel("Estimate of Probability", fontsize=16)
# fix the y-axis to be between 0 and 1:
ax.set_ylim(.4,.6)
# include a legend:
ax.legend()
# put a faded grid behind the graphic
ax.grid(True, alpha=0.25)

```

Notice that for very few flips the estimate of the probability is understandably poor.  But as the number of flips increases the estimate settles down to very close to the expected $0.5$. 

***

### Example 3- Drawing Cards

You randomly draw 2 cards (without replacement) from a standard 52-card deck.  What's the probability that you draw 2 hearts?


a).  Compute the probability by hand

b).  Write a simulation to verify your results.



**Simulation**:

```python
# Create a 52 card deck (use string 'D2' to represent the 2 of Diamonds)


suits=['D','H','C','S']
ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', '10', 'J', 'Q', 'K']


# Build a deck
deck = ...

print(deck)

```

```python

def check_two_hearts(twocards):   
    """ Function to simulate a single draw of 2 cards (WITHOUT REPLACEMENT)
        and check if both draws are hearts."""
    ...
    
    return ...
       
    
```

Test your function:

```python
check_two_hearts(['DA','H8']) 

#assert check_two_hearts(['H7','S7']) == False

```

Notice: In this scenario we want to draw two cards **WITHOUT REPLACEMENT**


**OPTION 1:** 
`rng.choice(cards, replace=False, size=2)`


**OPTION 2:** `rng.shuffle(cards)`  and then choose the first 2 in the shuffled list




 

```python
print(deck[:5])

rng.shuffle(deck)

print("Shuffled Deck", deck[:5])
```

```python

def single_draw(deck):

    """ Function that takes in a card deck, randomly draws 2 cards without replacement
    and checks if both are hearts. Returns a Boolean. """

    rng.shuffle(deck)
    first_two=deck[:2]
    #print(first_two)
    
    return check_two_hearts(first_two)
  

out = single_draw(deck)
out
```

```python
def probability_of_two_hearts(deck, num_sim):
    
    """Function that takes in a card deck and conducts num_sim simulations of drawing 2 cards (without replacement)
    and checking if they're both hearts. 
    Returns simulated probability out of num_sim simulations as well as a list with running probabilities."""
    
    ...
    
    return ...



```

This is very close to the actual answer we analytically derived in part (a).

```python
# Write code to plot the running probability:

x = np.arange(1, 10001)

prob, running_frac = probability_of_two_hearts(deck, num_sim=10000)


print("Simulated probability", prob)

# Plot running estimate of probability of getting two hearts as num_trials gets larger:
fig, ax = plt.subplots(figsize=(12,6))

# plot the terms in p.   Use matplotlib's plt.plot function


ax.plot(..., ..., color="steelblue",label="Simulated probability")

#Plot the theoretical probability
plt.axhline(y = ..., color = 'r', linestyle = '-', label = "Theoretical probability")

# put labels on the axes and give the graph a title.
ax.set_title("Running Estimate of Probability of Two Hearts", fontsize=20)
ax.set_xlabel("Number of Draws", fontsize=16)
ax.set_ylabel("Estimate of Probability", fontsize=16)
# fix the y-axis to be between 0 and 1:
ax.set_ylim(0,0.1)
# include a legend:
ax.legend()
# put a faded grid behind the graphic
ax.grid(True, alpha=0.25)
```

*** 

### Example 4- Conditional Probabilities with Dice


Suppose you roll a fair die two times.  Let $A$ be the event "the sum of the throws equals 4" and $B$ be the event "at least one of the throws is a $3$"

a).   Compute (by hand) the probability that the sum of the throws equals 4 _given_ that at least one of the throws is a 3.  That is, compute $P(A \mid B)$. 


**b).** Write code to simulate the the conditional probability $P(A \mid B)$. **Hint**: Think about the definition of conditional probability.

*Hint:  the Numpy methods `np.logical_or` and `np.logical_and` are potentially useful.*

```python
def simulateA_given_B(num_samples):
    """Write code that conducts num_samples simulations to estimate 
    the probability that when rolling two die, the sume of the throws equals 4 
    given that at least one of the throws is a 3.  Return the estimated probability"""
    
    ...
    
    
    
    
    return ...
                            
    
cond_prob = simulateA_given_B(1000)
    
print(cond_prob)                            
    
```

<!-- BEGIN QUESTION -->

***


<a id='p4'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 4a: Poker! (4 pts) ## 

A common example for discrete counting and probability questions are poker hands.  Consider using a standard 52-card playing deck, with card ranks `[A,2,3,4,5,6,7,8,9,10,J,Q,K]` across the standard 4 suits: `[C,D,H,S]`.

Suppose we draw 5 cards at random from the deck without replacement.  

In Poker, **"Three of a Kind"** is defined as a hand that contains **three cards of one rank and two cards of two other ranks.**  Notice that in this definition, a **Full House** (a hand that contains three cards of one rank and two cards of another rank) is **NOT classified as "three of a kind".** https://en.wikipedia.org/wiki/List_of_poker_hands#Three_of_a_kind


<img src="img/three.png" width="330px">

**What is the probability of drawing 5 cards (without replacement) that are "three of a kind?"**
   

Typeset your work using LaTeX below.  Show work justifying all steps.  You may leave your answer in terms of a ratio of products, but you should simplify away any combinatoric notation such as ${n \choose k}$ or $P(n,k)$.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for 4a, enter the answer you came up with (just the probability) in the cell below. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer in the cell above.

```python
q4a_answer = ...
```

```python
grader.check("q4acheck")
```





One of the big payoffs of simulation is that it can let us answer some probability questions that are otherwise quite difficult.  We can instead just simulate the process and get approximate results based on simulation.

To simulate the scenario you calculated theoretically in part 4a, you will use the steps below to write code to simulate draws of a hand of 5 cards from a standard 52-card deck.  You'll repeat the simulation 10000 times, and report how many "three of a kind" hands were observed.  If your simulation code is correct, your answer should be very close to your theoretical answer from part 4a. 

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 4bi (2 pts) ## 
    
Use the Python lists `suits` and `ranks` provided below to create a Python `list` called `cards` with 52 strings that represnt a full deck of cards. 

In the string describing the card give the suit first and the rank second.  For example to represent the 3 of hearts use the string 'H3'. To represent the queen of diamonds use the string 'DQ'.   

Use loops and/or list comprehensions to create the `card` list - **do not enter in the values manually.** 
Your deck should have 52 distinct cards.  


```python
suits=['D','H','C','S']

ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', '10', 'J', 'Q', 'K']

...
# 4bi).  Write code to create a list called `cards` representing a 52-card deck above this line.

cards
```

```python
grader.check("q4bi")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 4bii (3 pts) ## 

Define a function named `three_kind` to take a Python `list` of 5 cards as input (where the cards are given as strings such as `H3`) and check if the cards are Three of a Kind.  Your function should return `True` if the input is 3 of a kind and `False` otherwise. 

HINT: The function `np.unique` (https://numpy.org/doc/stable/reference/generated/numpy.unique.html) can make checking hands easier. 


```python
...
# 4bii).Write a function (above this line) 'three_kind' to check a list of 5 cards to see if it is 3 of a kind 
```

```python
grader.check("q4bii")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 4biii (4 pts) ## 

Write code in the space below that completes the following steps:

Step 1:  Write a function to simulate 10,000 random draws from `cards` of 5 cards each, and check if each draw is Three of a Kind. The function should return the overall proportion of random hands (out of the 10,000) in which Three of a Kind was observed. If you have coded your simulation correctly, your answer to this part should be very close to your theoretical answer from Part 4a.  

Step 2:  Let's visualize how this simulation converges to the theoretical probability.   In the video in Question 3, we plotted a running estimate of the probability of an event as a function of the number of trials in our simulation.  Write code that completes 10,000 random draws of 5 cards each, but this time outputs a plot of a running estimate of the proportion of hands that are Three of a Kind as a function of the number of trials (from 1 to 10,000) in your simulation.  **Include a red horizontal line on your plot with the theoretical probability that you calculated in part 4a.**  Be sure to include a title on your plot and be sure to label both your axes on the plot, as well as include a legend.






```python
...
#4biii).Write your code for Step 1 above this line
```

```python
...
# 4biii) Write your code for Step 2 above this line
```

<!-- END QUESTION -->



<a id='p5'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 5:  (8 pts) ## 

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 5a (4 pts) ## 

**5a i).**  What is the **mathematical definition** of the conditional probability of an event A given an event B?  Give your answer using LaTex and correct mathematical notation.   You can assume $P(A)>0$ and $P(B)>0$


**5a ii).** You flip a fair coin twice.  What is the probability that both flips are heads given that at least one of the flips is a heads?  Explain/justify using math.


**Enter your answer for part 5a i)** in this cell (double click on this cell and write all steps using Markdown and LaTex):

**Enter your answer for part 5a ii)** in this cell (double click on this cell and write all steps using Markdown and LaTex):

<!-- END QUESTION -->

**Q5aii Final Answer Check).**  To check your final answer to Q5aii, enter the final answer you came up with **as a fraction** in the cell below. Note that this is just a built-in public test so you can check your work and determine if you are on the right track.  To receive credit on this problem you must show all steps above using LaTeX and fully justifying your answer using correct mathematical notation.  

```python
q5aii_answer = ...
```

```python
grader.check("q5aiicheck")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />


## 🔴  Question 5b (4 pts):

**5b).**  Suppose a bag has 5 red marbles and 3 blue ones.  We select 2 marbles without replacement (i.e. we select the first marble and then we don't put it back and then select a 2nd marble). 
    
 - i).  What is the probability that the 2nd marble is red given that the first marble is red?  Explain/justify using math. 
    
 - ii).  What is the probability that both marbles are red?  Explain/justify using math.  

**Enter your answer for part 5bi)** in this cell (double click on this cell and write all steps using Markdown and LaTex):

**Enter your answer for part 5bii)** in this cell (double click on this cell and write all steps using Markdown and LaTex):

<!-- END QUESTION -->

**5b Final Answer Check).**  To check your final answers to 5b, enter the final answer you came up with (**as simplified fractions**) in the cell below. Note that this is just a built-in public test so you can check your work and determine if you are on the right track.  To receive credit on this problem you must show all steps in part 5b above using LaTeX and fully justifying your answer using correct mathematical notation.  

```python
q5bi_answer = ...

q5bii_answer = ...
```

```python
grader.check("q5bcheck")
```

<!-- BEGIN QUESTION -->



<a id='p6'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 6 (7 pts) ## 

The accuracy of a diagnostic test is often described using the following terms:

 - Test Sensitivity: Ability to detect a positive case (i.e. probability that the test is positive given that the person actually has the virus).

 - Test Specificity:  Ability to determine a negative case (i.e. the probability that a person tests negative given that they don't have the virus)).


Suppose a  diagnostic test for a virus is reported to have $85\%$ sensitivity and $95\%$ specificity.

Suppose $3\%$ of the population has the virus in question.  

Answer the following questions all in ONE cell below using LaTeX.  Show all steps.

**6a) (6 pts)**.  If a person is chosen at random from the population and the diagnostic test indicates that they have the virus, what is the conditional probability that they do, in fact, have the virus?    Write up your full solution using LaTeX.

**6b) (1 pt)**.  Terminology: What is the **prior** and what is the **likelihood** in this scenario?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for Problem 6a, enter the answer you came up with (just the probability) in the cell below.  Do not round your answer - give the exact probability. Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer above.

```python
q6a_answer = ...
```

```python
grader.check("q6check")
```

<!-- BEGIN QUESTION -->

***


<a id='p7'></a>

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 7a (3 pts) ## 

To play a game, you have a bag containing 30 fair [four-sided dice](https://en.wikipedia.org/wiki/Four-sided_die), with faces `{1,2,3,4}`.  This bag also contains 9 fair six-sided dice (faces `{1,2,3,4,5,6}`) and 3 fair twenty-sided dice (faces `{1,2,3,4,...,19,20}`).  Call these 3 classes of die "Four", "Six" and "Twenty" (or $D_4$, $D_6$, and $D_{20}$, for short).  You grab one die at random from the box.  

Work the following problems by hand and write up your full solution using LaTeX unless otherwise stated (but don't be afraid to simulate to check your result!). 

You grab one die at random from the box and roll it one time.  What is the probability of the event $R_2$, that you roll a 2? Explain your reasoning mathematically (using LaTeX).   

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
# Optional space for simulation/checking work:


```

To check your answer for 7a, enter the answer you came up with (just the probability) in the cell below.  Do not round your answer- give the exact probability.  Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer in part 7a above.

```python
q7a_answer = ...
```

```python
grader.check("q7acheck")
```

<!-- BEGIN QUESTION -->

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 🔴   Question 7b (3 pts) ## 

You grab one die at random from the box and roll it one time. Suppose you roll a 2.  Given this information, what is the probability that the die you chose from the box is a Twenty-sided die? Write up your full solution using LaTeX. Show all steps.  

_Type your answer here, replacing this text._

<!-- END QUESTION -->

To check your answer for 7b, enter the answer you came up with (just the probability) in the cell below. Give the exact probability (do not round).  Note that this is just a public test. Full credit requires that you show all your work written in LaTeX and fully justifying your answer in part 7b above.

```python
q7b_answer = ...
```

```python
grader.check("q7bcheck")
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 4!

If you discussed this assignment with any other students in the class (in a manner that is acceptable as described by the Collaboration policy above) please **include their names** here:

**Collaborators**: *list collaborators here*

If while completing this assignment you reference any websites other than those linked in this assignment or provided on Canvas please list those references here:

**External references**:  *list any external websites you referenced*

```python
# NOTEBOOK INDENTATION CHECK
# This checks for any extra spaces you accidently created that will cause your notebook to fail tests when you upload to Gradescope. 
# Save your notebook before running this cell.

import json
from pathlib import Path

MARKER = "NOTEBOOK_INDENTATION_CHECK"

# Find the saved notebook containing this code
notebooks = []

for path in Path.cwd().glob("*.ipynb"):
    try:
        if MARKER in path.read_text(encoding="utf-8"):
            notebooks.append(path)
    except Exception:
        pass

if not notebooks:
    print("ERROR: Could not find the saved notebook.")
    print("Save your notebook and run this cell again.")

else:
    # If more than one copy exists, use the most recently saved one
    notebook_file = max(notebooks, key=lambda p: p.stat().st_mtime)

    print("Checking:", notebook_file.name)
    print()

    with open(notebook_file, "r", encoding="utf-8") as f:
        nb = json.load(f)

    problems = []

    for cell_number, cell in enumerate(nb["cells"], start=1):

        if cell.get("cell_type") != "code":
            continue

        source = "".join(cell.get("source", []))
        lines = source.splitlines()

        try:
            compile(source, "cell", "exec")

        except (IndentationError, TabError) as e:
            bad_line = ""

            if e.lineno and 1 <= e.lineno <= len(lines):
                bad_line = lines[e.lineno - 1]

            problems.append(
                (cell_number, e.lineno, e.msg, bad_line)
            )

        except SyntaxError:
            pass

    if not problems:
        print("PASS: No indentation problems were found.")

    else:
        print("WARNING: Possible indentation problem(s) found.")
        print()

        for cell_number, line_number, message, bad_line in problems:

            print("-" * 60)
            print("Notebook cell:", cell_number)
            print("Line in cell:", line_number)
            print("Problem:", message)
            print()
            print("Problem line:")
            print(repr(bad_line))

            search_text = bad_line.strip()

            if search_text:
                print()
                print("Use Cmd+F or Ctrl+F and search for:")
                print(search_text[:100])

            print()

        print("Fix the problem(s), save the notebook,")
        print("and run this check again.")
```

### Submission Instructions

Before proceeding any further, **save this notebook.**

After running the `grader.export()` cell provided below, **2 files will be created**: a zip file and pdf file.  You can download them using the links provided below OR by finding them in the same folder where this juptyer notebook resides in your JuptyerHub.

To receive credit on this assignment, **you must submit BOTH of these files
to their respective Gradescope portals:** 

* **Homework 4 Autograded**: Submit the zip file that is output by the `grader.export()` cell below to the HW4 Autograded assignment in Gradescope.

* **Homework 4 Manually Graded**: Submit your hw04.PDF to the HW4 Manually Graded assignment in Gradescope.  **YOU MUST SELECT THE PAGES CORRESPONDING TO EACH QUESTION WHEN YOU UPLOAD TO GRADESCOPE.  IF NOT, YOU WILL LOSE POINTS**   Also, **check** that all of your plots **and** all lines of your code are showing up in your PDF before submitting.  If not, you will not receive credit for your plots/code.  


**You are responsible for ensuring your submission follows our requirements. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

**TROUBLESHOOTING TIPS**  If you are having any issues compiling your assignment, please [read through these troubleshooting tips first](https://docs.google.com/document/d/1ndr3Wj1PSF5qzlLMaBJznwh6QGeEXjd5TAJ6nf9EJvo/edit?usp=sharing), then post any questions on Piazza. 

```python
import simple_latex_checker as slc

nb = slc.Nb_checker()
nb.run_check("hw04.ipynb")
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

AFTER running the cell below, click on <a href='hw04.pdf' download>this link to download the PDF </a> to upload to Gradescope.  There will be a separate link that appears after running the cell below with a link to download the zip file to upload to Gradescope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1a (points: 1)
```
>>> assert get_hash(round(first_three_black, 4)) == '271991f15c06e3dd9aa7a0d196ff5b19'
```


Tests for q1b (points: 1)
```
>>> assert get_hash(round(no_green, 4)) == '157efb8545e3fd05e47afd97f4816d17'
```


Tests for q1c (points: 1)
```
>>> assert get_hash(round(at_least_one_green, 4)) == 'bdfd2bbc94051b89092784727dadb9c9'
```


Tests for q1d (points: 1)
```
>>> assert get_hash(round(lone_winners, 4)) == '4ad22987066120a9962f994a9fe684f4'
```


Tests for q2acheck (points: 0)
```
>>> assert get_hash(round(q2a_answer, 4)) == '3d522deaf85577451c01974654b36ad3'
```


Tests for q2bcheck (points: 0)
```
>>> assert get_hash(round(q2b_answer, 4)) == 'f452927b0b2a755192117d956a8d47f6'
```


Tests for q4acheck (points: 0)
```
>>> assert get_hash(round(q4a_answer, 4)) == 'ace990a0524b7549fea9afe4a158c01c'
```


Tests for q4bi (points: 2)
```
>>> assert get_hash(len(cards)) == '9a1158154dfa42caddbd0694a4e9bdc8'
```
```
>>> assert get_hash(sorted(set(cards))) == '83dfc5b12f2790365c2b2ba676e90de0'
```


Tests for q4bii (points: 3)
```
>>> assert three_kind(['H3', 'S3', 'D3', 'H4', 'C5']) == True
```
```
>>> assert three_kind(['C4', 'D4', 'S4', 'H4', 'H2']) == False
```
```
>>> assert three_kind(['SA', 'S4', 'S4', 'H5', 'H10']) == False
```
```
>>> assert three_kind(['SA', 'DA', 'HA', 'C5', 'H5']) == False
```
```
>>> assert three_kind(['S3', 'D3', 'H4', 'C4', 'H5']) == False
```
```
>>> assert three_kind(['S10', 'D10', 'H10', 'CQ', 'HA']) == True
```


Tests for q5aiicheck (points: 0)
```
>>> assert get_hash(round(q5aii_answer, 4)) == 'd80cd8a9bc59ca0c1f479e4eca3c64be'
```


Tests for q5bcheck (points: 0)
```
>>> assert get_hash(round(q5bi_answer, 4)) == '790abc5c38e7c740420b03c24fabb05b'
```
```
>>> assert get_hash(round(q5bii_answer, 4)) == 'b4b8f49579f52a53c6963b8f8f685a29'
```


Tests for q6check (points: 0)
```
>>> assert get_hash(round(q6a_answer, 5)) == '439546d788c53a6c0c1c5194027e5d36'
```


Tests for q7acheck (points: 0)
```
>>> assert get_hash(round(q7a_answer, 6)) == 'cb96698ec25a583b061cc9b57a4d3e36'
```


Tests for q7bcheck (points: 0)
```
>>> assert get_hash(round(q7b_answer, 4)) == 'c29ea1fe96eb4921715b8ceae933a374'
```
